# **🚀 Базовая линейная алгебра в Julia**  
# **Урок Двенадцатый.**

**Автор**: Андреас Ноак Йенсен (MIT & JuliaComputing) [(Twitter)](https://twitter.com/anoackjensen?lang=en)  
(с правками от Джейн Херриман и Сергея Соболевского)

<br>

##  **📌 Темы:**  

1. **Основные понятия линейной алгебры**
2. **Создание матрицы и вектора в Julia.**  
3. **Стандартные матричные операции в Julia.**
4. **Решение линейных систем (Solving Linear Systems).** 
5. **Библиотека `LinearAlgebra`.**



<br>

---

 <br>


# **Базовая линейная алгебра в Julia.**

**Линейная алгебра** — это раздел математики, изучающий векторы, матрицы, линейные уравнения и линейные преобразования. Она является основой для многих областей, включая компьютерные науки, машинное обучение, обработку сигналов и физику.

<br>

В **Julia** для вычисления определителя матриц больших размерностей ($ n>3 $) используется функция `LinearAlgebra.det` из стандартной библиотеки `LinearAlgebra`,
**Julia** использует LU-разложение (аналогично методу Гаусса), что обеспечивает высокую эффективность даже для очень больших матриц ($n≥1000$). Для работы с символьными матрицами в Julia
можно подключить пакет `Symbolics.jl`, для численных вычислений используется тип данных матрицы (например, `Float64` или `BigFloat` для повышенной точности).

<br>

---

<br>

### **1️⃣ Создание матрицы и вектора в Julia.**  

В Julia векторы и матрицы представляют собой массивы `(Array)`, но они имеют разную размерность. 

**Вектор** — это одномерный массив, а **Матрица** — двумерный.

#### **🔹 Определение заданной матрицы** 

Сначала создадим простую иатрицу А размером 3х3:

```julia
A = [1 2 3; 
     4 5 6; 
     7 8 9]  # Задаем матрицу A (3×3)
```

```julia
A = [1 2 3; 4 5 6; 7 8 9]
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 4  5  6
 7  8  9
```


#### **🔹 Определение случайной матрицы**  

Теперь создадим **случайную матрицу** `A` размером **3×3**, содержащую целые числа от 1 до 10:  


```julia
A = rand(1:10, 3, 3) # Определяем матрицу 3x3, заполненную случайными числами от 1 до 10

```

Output:
```
3×3 Matrix{Int64}:
 3  3  7
 4  5  7
 7  8  1
```

#### **🔹 Определение вектора из единиц.** 
 
Теперь создадим **вектор**, содержащий только целые числа - единицы `1`, размером **3 элемента**:  

```julia
x = fill(1,3)  # Определяем вектор x = [1, 1, 1]
```

Output:
```
3-element Vector{Int64}:
 1
 1
 1
```

#### **🔹 Определение вектора случайных чисел.** 

С помощью функциии `rand()` определяем вектор `b`, заполненный случайными числами от 1 до 10 и размерностью 5 элементов:

```julia
b = rand(1:10, 5)    # Определяем вектор b длины 5, заполненный случайными числами от 1 до 10
```

Output:
```
5-element Vector{Int64}:
 2
 9
 3
 6
 8
```

#### **🔹 Определение вектора случайных чисел от 0 до 1.**

Функция `rand(n)` в Julia по умолчанию генерирует **n** случайных чисел из равномерного распределения на интервале **[0, 1]**.

```julia
z = rand(3) # вектор x длины 3, заполненный случайными числами от 0 до 1
```

Output:
```
3-element Vector{Float64}:
 0.09639256229404969
 0.5434733985340262
 0.027005953640411584
```

#### **🔹 Функция collect() в Julia.**

Основное использование `collect()` это преобразование диапазона `(Range)` в массив `(Vector)`.

В Julia диапазон `(1:5)` — это ленивый объект, который хранит только начальное и конечное значение, но не создает массив.
Если вам нужен массив значений диапазона, используйте `collect()`:

```julia
x = 1:5       # Диапазон от 1 до 5
v = collect(x)  # Преобразование в массив (вектор)

```

Output:
```
5-element Vector{Int64}:
 1
 2
 3
 4
 5
```

```julia
k = collect(1:3)   # Вектор из диапазона от 1 до 3
```

Output:
```
3-element Vector{Int64}:
 1
 2
 3
```

#### **🔹 Определение вектора из чисел с плавающей запятой.**


Для примера создадим вектор с длинной `3` и заполненный `1.55` - числами с плавающей точкой (тип Float64):

```julia
y = fill(1.55, 3) # вектор x длины 3, заполненный 1.55 - числами с плавающей точкой
```

Output:
```
3-element Vector{Float64}:
 1.55
 1.55
 1.55
```

#### **🔹 Подведём итог по полученным типам данных в Julia.**  

В Julia **векторы и матрицы** представляют собой **разные структуры данных**, имеющие разные размеры и типы:  

- **Матрица** `A` имеет тип **`Matrix{Int64}`** (`Array{Int64,2}`) – это **двумерный массив**.  
  
- **Вектор** `x` имеет тип **`Vector{Int64}`** (`Array{Int64,1}`) – это **одномерный массив** с целыми числами. 
   
- **Вектор** `y` имеет тип **`Vector{Float64}`** (`Array{Float64,1}`) – это **одномерный массив** с числами с плавающей запятой.  

<br>

##### **📌 Важные псевдонимы в Julia.**

Для удобства в Julia используются **псевдонимы типов массивов**: 

- **`Vector{Type} ≡ Array{Type,1}`** → одномерный массив (вектор).
    
- **`Matrix{Type} ≡ Array{Type,2}`** → двумерный массив (матрица).  

Это делает код **более читаемым и удобным**, особенно при работе с линейной алгеброй.  

<br>

##### **📌 Основные выводы:**

✅ Векторы (`Vector{T}`) – **одномерные** массивы, матрицы (`Matrix{T}`) – **двумерные**.

✅ **Операции с матрицами и векторами** в Julia работают **так же, как в других языках**, но более **оптимизированы**. 

✅ **Julia автоматически выбирает подходящий тип данных**, обеспечивая **высокую производительность**.  



<br>

---

<br>

### **2️⃣ Стандартные матричные операции в Julia.**

<br>

### **🔹 1. Умножение (Multiplication)**  
Julia поддерживает стандартные **матричные операции**, такие как умножение матрицы на вектор.  

```julia
b = A * x
```

Output:
```
3-element Vector{Float64}:
 9.0
 8.0
 6.0
```

📌 Оператор `*` используется для **матричного умножения** (`A` на `x`).  

🔹 **Результат** будет вектором `Vector{Float64}`, так как `A` — целочисленная (`Int64`), а `x` содержит числа с плавающей запятой (`Float64`).

<br>

### **🔹 2. Транспонирование и сопряженное транспонирование (адъюнгирование) матрицы**  

В Julia, как и в математике, различают два вида транспонирования:  

1. **Обычное транспонирование** – перестановка строк и столбцов.  
   
2. **Сопряженное транспонирование (адъюнгирование)** – транспонирование с одновременным комплексным сопряжением элементов.  


#####  Транспонирование — это операция в линейной алгебре, при которой строки матрицы преобразуются в столбцы, и наоборот. Результатом транспонирования матрицы $ A $ является новая матрица $ A^T $ (или $ A'$), где строки исходной матрицы становятся столбцами, а столбцы — строками.

<br>

### Пример:

Рассмотрим матрицу $ A $ размером $ 2 \times 3 $:

$$ A = \begin{pmatrix}
1 & 2 & 3 \\
4 & 5 & 6
\end{pmatrix} $$

Транспонированная матрица $ A^T $ будет иметь размер $ 3 \times 2 $:

$$ A^T = \begin{pmatrix}
1 & 4 \\
2 & 5 \\
3 & 6
\end{pmatrix} $$

### Свойства транспонирования:

1. **Двойное транспонирование**: Транспонирование транспонированной матрицы возвращает исходную матрицу:
   $ (A^T)^T = A $

2. **Сумма матриц**: Транспонирование суммы матриц равно сумме транспонированных матриц:
   $ (A + B)^T = A^T + B^T $

3. **Произведение матриц**: Транспонирование произведения матриц равно произведению транспонированных матриц в обратном порядке:
   $ (AB)^T = B^T A^T $

4. **Скалярное произведение**: Транспонирование произведения матрицы на скаляр равно произведению транспонированной матрицы на тот же скаляр:
   $ (cA)^T = cA^T $

Транспонирование широко используется в различных областях математики, физики, инженерии и компьютерных наук для упрощения вычислений и решения задач, связанных с матрицами.

<br>

#### **2.1 Обычное транспонирование**

Обычное транспонирование просто меняет строки и столбцы местами.

**Например:**

Создадим матрицу `А`:

```julia
A = [1 2 3; 
     4 5 6; 
     7 8 9]  # Задаем матрицу A (3×3)
```

```julia
A = [1 2 3; 4 5 6; 7 8 9]
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 4  5  6
 7  8  9
```

Теперь её транспонируем:

```julia
transpose(A)
```

Output:
```
3×3 transpose(::Matrix{Int64}) with eltype Int64:
 1  4  7
 2  5  8
 3  6  9
```

🔹 **Результат**:

Строки и столбцы поменялись местами

```julia
3×3 Matrix{Int64}:
 1  4  7
 2  5  8
 3  6  9
 
```
📌 **Обычное транспонирование используется для любых матриц и не затрагивает значения элементов.**  

<br>

#### **2.2 Сопряженное транспонирование (адъюнгирование)**  

Если матрица содержит **комплексные числа**, то операция `A'` выполняет **сопряженное транспонирование**.

Наша матрица `А` с комплексными числами:

```julia
A = [1+2im 3-im; 4 5+3im]
```

Output:
```
2×2 Matrix{Complex{Int64}}:
 1+2im  3-1im
 4+0im  5+3im
```

### **🔹 2.
Теперь выполним сопряженное транспонирование:

```julia
A'
```

Output:
```
2×2 adjoint(::Matrix{Complex{Int64}}) with eltype Complex{Int64}:
 1-2im  4+0im
 3+1im  5-3im
```

🔹 **Результат**:
 
```julia
2×2 Matrix{Complex{Int64}}:
 1 - 2im   4
 3 + 1im   5 - 3im
 
```
📌 **Что мы видим?** 

- Что матрица **Транспонируется** , т.е. поменялись местами строки и столбцы.  
  
- Что произошло **Комплексное сопряжение**, т.е. у элементов матрицы `А` изменился знак мнимой части \( i \) на противоположный.  


👉 **Если матрица содержит только вещественные числа, `A'` работает так же, как `transpose(A)`.**  


<br>

#### **🔹 3. Транспонированное умножение**  

В Julia можно записывать **умножение транспонированной матрицы на саму себя** без оператора `*` как `A'A`:  



```julia
A = [1 2 3; 4 5 6; 7 8 9] # задаем матрицу А

A'A # умножаем матрицу на транспонированную матрицу
```

Output:
```
3×3 Matrix{Int64}:
 66   78   90
 78   93  108
 90  108  126
```

`A'A` - это сокращенная запись для `transpose(A) * A`, проверим:

```julia
A = [1 2 3; 4 5 6; 7 8 9]

transpose(A) * A
```

Output:
```
3×3 Matrix{Int64}:
 66   78   90
 78   93  108
 90  108  126
```

📌 **Это используется в линейной алгебре**, например, для вычисления **граммианов** и **норм уравнений в методе наименьших квадратов**.  

<br>

---

<br>

### **3️⃣ Решение линейных систем (Solving Linear Systems)**  

<br>

Оператор **обратного деления `\`** - решает линейные уравнения вида **$Ax = b$**. В зависимости от формы матрицы $ A $ **Julia** автоматически выбирает подходящий метод решения.
<br>

#### **⬜ 1. Квадратная система (Square System)**  

Если \( A \) – **квадратная матрица** (число строк = числу столбцов), оператор `\` находит **точное решение**:  

```julia
A = rand(3, 3)
b = rand(3)

x = A \ b  # Решает Ax = b
```

Output:
```
3-element Vector{Float64}:
  0.027035062177010568
  1.5815023677620113
 -0.5161651032603933
```

📌 **Если матрица невырождена**, Julia использует метод **LU-разложения** для быстрого решения.

<br>

#### **⬜ 2. Переопределенная система (Overdetermined, "высокая" матрица)**  

Если уравнение имеет **больше уравнений, чем переменных** (**"высокая" матрица**), Julia находит **наилучшее среднеквадратичное решение** (метод наименьших квадратов):  


```julia
Atall = rand(3, 2)  # 3 строки, 2 столбца (больше уравнений, чем переменных)
b = rand(3)

x = Atall \ b  # Решает Ax = b
```

Output:
```
2-element Vector{Float64}:
  1.106749415030229
 -0.279873599951611
```

📌 **Julia автоматически использует метод QR-разложения**.

<br>

#### **⬜ 3. Недоопределенная система (Underdetermined, "короткая" матрица)**  

Если уравнение имеет **меньше уравнений, чем переменных** (**"короткая" матрица**), Julia вычисляет **минимальное нормированное решение**:


```julia
Ashort = rand(2, 3)  # 2 строки, 3 столбца (меньше уравнений, чем переменных)
bshort = rand(2)

x = Ashort \ bshort  # Решает Ax = b
```

Output:
```
3-element Vector{Float64}:
 0.4419893495217756
 0.4105939591388945
 0.6770233980348624
```

📌 **Julia использует SVD-разложение**, гарантируя **минимальную норму** решения.

<br>

#### **⬜ 4. Неполный ранг (Rank-Deficient System)**  

Если \( A \) содержит **зависимые столбцы** (линейно зависимые строки), то Julia решает уравнение с **минимальной нормой**:  

```julia
v = rand(3)
rankdef = hcat(v, v)  # Создаем матрицу с зависимыми столбцами

x = rankdef \ b  # Найдет минимальное нормированное решение
```

Output:
```
2-element Vector{Float64}:
 0.506204980481362
 0.5062049804813621
```

📌 **Julia использует псевдообратную матрицу (Moore-Penrose inverse) через SVD.**

#### **😺 Важно запомнить:** 

✅ **`transpose(A)`** – обычное транспонирование.

✅ **`A'`** – транспонирование + комплексное сопряжение. 

✅ **`A'A`** – транспонированное умножение.

✅ **`A \ b`** – оператор решения линейных систем, автоматически выбирает оптимальный метод.  

**💡 Julia делает линейную алгебру интуитивной и мощной, автоматизируя сложные вычисления!** 🚀

<br>

---

<br>

### **4️⃣  Библиотека `LinearAlgebra` в Julia.**

Хотя **основные функции линейной алгебры** доступны в Julia **по умолчанию**, стандартная библиотека **`LinearAlgebra`** предоставляет **дополнительные инструменты** для работы с матрицами и векторами.  

📌 **Что добавляет `LinearAlgebra`?**  
- **Факторизации** (LU, QR, Cholesky, SVD и др.).  
- **Структурированные матрицы** (диагональные, разреженные, симметричные и др.).  
- **Вычисления норм, определителей, ранга, следа и др.**  
- **Операции с векторами** (скалярное, векторное, смешанное произведение).  

📌 **Как подключить библиотеку?**  

```julia
using LinearAlgebra  # Подключаем библиотеку LinearAlgebra
```

После этого **новые функции** становятся доступными в текущем сеансе. 

<br>

### **🔹 Основные возможности `LinearAlgebra`**

##### ✅ **1. Нормы и ортогональность**  
Функция `norm()` вычисляет **разные нормы** векторов и матриц:  

```julia
using LinearAlgebra

x = [3, 4]
println(norm(x))       # Евклидова норма (по умолчанию)
println(norm(x, 1))    # 1-норма (сумма модулей элементов)
println(norm(x, Inf))  # Бесконечная норма (максимальный элемент)
```

##### ✅ **2. Определитель матрицы**  

```julia
A = [1 2; 3 4]
det(A)  # -2.0
```

##### ✅ **3. Ранг и след матрицы**  

```julia
rank(A)   # 2 (ранг матрицы)
tr(A)     # 5 (след матрицы, сумма диагональных элементов)
```

##### ✅ **4. Факторизации (разложения матриц)**  

```julia
A = [4 3; 6 3]
F = lu(A)    # LU-разложение
println(F.L) # Нижняя треугольная матрица
println(F.U) # Верхняя треугольная матрица
```

Также доступны `qr(A)`, `svd(A)`, `cholesky(A)`, `eigen(A)`.

##### ✅ **5. Работа с особыми матрицами**  

```julia
D = Diagonal([1, 2, 3])   # Диагональная матрица
S = Symmetric(A)          # Симметричная матрица
```

##### **🔹 Подводим итог по `LinearAlgebra`:**

✅ **`LinearAlgebra` расширяет стандартные функции Julia для работы с матрицами и векторами**.

✅ Включает **факторизации, вычисление норм, работу с особыми матрицами**.  

✅ Подключается командой **`using LinearAlgebra`**.  

**💡 Важно!** 

**Если вы работаете с линейной алгеброй в Julia, то этот пакет для вас - незаменимый инструмент!**

<br>

---

<br>

### **Упражнения** 🚀 

<br>

#### ✅ Здание 11.1  

Возьмите скалярное произведение (или «скалярное» произведение) вектора `v` на самого себя и присвойте его переменной `dot_v`.



```julia
v = [1,2,3]
```

Output:
```
3-element Vector{Int64}:
 1
 2
 3
```

```julia
# Ваше решение:

```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

<br>

#### ✅ Здание 11.2 

Возьмите внешнее произведение вектора v на себя и присвойте его переменной `outer_v`

```julia
# Ваше решение:

```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

<br>

#### ✅ Здание 11.3 

Используйте [Linear Algebra.cross](https://docs.julialang.org/en/v1/stdlib/LinearAlgebra/#LinearAlgebra.cross), чтобы вычислить взаимное произведение вектора v с самим собой и присвоить его переменной `cross_v`

```julia
# Ваше решение:
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```