# **🚀 Строки в Julia. Урок Второй.**

<br>

## **📌 Темы:**

1. **Как получить строку.**
   
2. **Интерполяция строк.**
   
3. **Объединение строк.**

<br>

---

<br>

### **1️⃣ Как получить строку.**

Заключите свои символы в " " или """" """ !

```julia
s1 = "Я сильный человек. Я сильно люблю программирование."
```

Output:
```
"Я сильный человек. Я сильно люблю программирование."
```

```julia
s2 = """Я тоже сильный. Я тоже люблю программирование."""
```

Output:
```
"Я тоже сильный. Я тоже люблю программирование."
```

📌 Существует несколько функциональных различий между строками, заключенными в одинарные (" ") и тройные кавычки(""" """).

Одно из отличий заключается в том, что в последнем случае вы можете использовать кавычки внутри вашей строки.

```julia
"Здесь мы получаем "ошибку", потому что неясно, где заканчивается эта строка"
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[3]:1
```

```julia
"""Смотри, мама, никаких "ошибок"!!! """
```

Output:
```
"Смотри, мама, никаких \"ошибок\"!!! "
```

Обратите внимание, что ```' '``` определяет символ, но НЕ строку!

```julia
typeof('a')
```

Output:
```
Char
```

```julia
'Здесь мы получим сообщение об ошибке'
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[6]:1
```

<br>

---

<br>


### **2️⃣ Интерполяция строк (String interpolation).**

Мы можем использовать знак  **$**  для вставки существующих переменных в строку и для вычисления выражения внутри строки. <br><br>
Ниже приведен пример, содержащий некоторую особо важную личную :)) информацию.

```julia
name = "Jane"
num_fingers = 10
num_toes = 10
```

Output:
```
10
```

```julia
println("Здравствуйте, меня зовут $name.")
println("У меня $num_fingers пальцев на руках и $num_toes пальцев на ногах.")
```

Output:
```
Здравствуйте, меня зовут Jane.
У меня 10 пальцев на руках и 10 пальцев на ногах.
```

```julia
 println("Это $(num_fingers + num_toes) пальцев всех вместе!!")
```

Output:
```
Это 20 пальцев всех вместе!!
```

<br>

---

<br>

### **3️⃣ Объединение строк (String concatenation).**

Ниже приведены три способа объединения строк! <br><br>
Первый способ - использовать функцию `string()`. <br>
`string()` преобразует нестроковые входные данные в строки.

```julia
s3 = "Сколько кошек, ";
s4 = "неужели их слишком много?";
😺 = 10
```

Output:
```
10
```

```julia
string(s3, s4)
```

Output:
```
"Сколько кошек, неужели их слишком много?"
```

```julia
string("Я не знаю, но ", 😺, " - это слишком мало.")
```

Output:
```
"Я не знаю, но 10 - это слишком мало."
```

Мы также можем использовать `*` для конкатенации!

```julia
s3*s4
```

Output:
```
"Сколько кошек, неужели их слишком много?"
```

<br>

---

<br>

## Упражнения 🚀

####  ✅  Задание 2.1 

Создайте строку, в которой "hi" повторяется 1000 раз, сначала с помощью "repeat", а затем с помощью оператора возведения в степень, который может вызывать "*". <br>
Присвойте ей переменную "hi", указанную ниже.

```julia
# Ваше решение:


```

<!-- grade_id: cell-5b38bdb1115161b4, points: 0, grade: false, solution: true -->
```julia
# Проверка:

# Использование функции repeat
hi1 = repeat("hi", 1000)

# Использование оператора возведения в степень (^)
hi2 = "hi"^1000

# Проверка на равенство
@assert hi1 == hi2

# Присвоение переменной hi
hi = hi1  # или hi = hi2, так как они идентичны

```

Output:
```
"hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"[93m[1m ⋯ 1460 bytes ⋯ [22m[39m"hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"
```

<!-- grade_id: cell-20b1027a968fbb48, points: 1, grade: true, solution: false -->
```julia
@assert hi == "hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"
println(hi)
```

Output:
```
hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi
```

<br>

####  ✅ Задание 2.2 

Объявите две переменные

```julia
a = 3
b = 4
```
и используйте их для создания двух строк:
```julia
"3 + 4"
"7"
```
и сохраните результаты в `c` и `d` соответственно

```julia
# Ваше решение:


```

<!-- grade_id: cell-36beb524c8ecd33b, points: 0, grade: false, solution: true -->
```julia
# Проверка

# Объявление переменных
a = 3
b = 4

# Создание строк
c = "$a + $b"   # строка "3 + 4"
d = "$(a + b)"  # строка "7"

# Проверка корректности
@assert c == "3 + 4"
@assert d == "7"

println(c)
println(d)

```

Output:
```
3 + 4
7
```

<!-- grade_id: cell-14c3c5b015a97892, points: 1, grade: true, solution: false -->
```julia
@assert c == "3 + 4"
@assert d == "7"

println(c)
println(d)
```

Output:
```
3 + 4
7
```