# **🚀 Строки в Julia. Урок Второй.**

<br>

## **📌 Темы:**

1. **Как получить строку.**
   
2. **Интерполяция строк.**
   
3. **Объединение строк.**

<br>

---

<br>

## 🎯 После этого занятия вы сможете

- создавать строки и понимать разницу между `String` и `Char`;
- пользоваться интерполяцией `$` и конкатенацией;
- объяснять, почему индексация строк в Julia работает по байтам, а не по символам;
- безопасно обходить строку с помощью `eachindex`, `nextind`, `firstindex`/`lastindex`;
- применять основные строковые функции и работать с Unicode.

### **1️⃣ Как получить строку.**

Заключите свои символы в " " или """" """ !

```julia
s1 = "Я сильный человек. Я сильно люблю программирование."
```

Output:
```
"Я сильный человек. Я сильно люблю программирование."
```

```julia
s2 = """Я тоже сильный. Я тоже люблю программирование."""
```

Output:
```
"Я тоже сильный. Я тоже люблю программирование."
```

📌 Существует несколько функциональных различий между строками, заключенными в одинарные (" ") и тройные кавычки(""" """).

Одно из отличий заключается в том, что в последнем случае вы можете использовать кавычки внутри вашей строки.

```julia
"Здесь мы получаем "ошибку", потому что неясно, где заканчивается эта строка"
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[3]:1
```

```julia
"""Смотри, мама, никаких "ошибок"!!! """
```

Output:
```
"Смотри, мама, никаких \"ошибок\"!!! "
```

Обратите внимание, что ```' '``` определяет символ, но НЕ строку!

```julia
typeof('a')
```

Output:
```
Char
```

```julia
'Здесь мы получим сообщение об ошибке'
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[6]:1
```

<br>

---

<br>


### **2️⃣ Интерполяция строк (String interpolation).**

Мы можем использовать знак  **$**  для вставки существующих переменных в строку и для вычисления выражения внутри строки. <br><br>
Ниже приведен пример, содержащий некоторую особо важную личную :)) информацию.

```julia
name = "Jane"
num_fingers = 10
num_toes = 10
```

Output:
```
10
```

```julia
println("Здравствуйте, меня зовут $name.")
println("У меня $num_fingers пальцев на руках и $num_toes пальцев на ногах.")
```

Output:
```
Здравствуйте, меня зовут Jane.
У меня 10 пальцев на руках и 10 пальцев на ногах.
```

```julia
 println("Это $(num_fingers + num_toes) пальцев всех вместе!!")
```

Output:
```
Это 20 пальцев всех вместе!!
```

<br>

---

<br>

### **3️⃣ Объединение строк (String concatenation).**

Ниже приведены три способа объединения строк! <br><br>
Первый способ - использовать функцию `string()`. <br>
`string()` преобразует нестроковые входные данные в строки.

```julia
s3 = "Сколько кошек, ";
s4 = "неужели их слишком много?";
😺 = 10
```

Output:
```
10
```

```julia
string(s3, s4)
```

Output:
```
"Сколько кошек, неужели их слишком много?"
```

```julia
string("Я не знаю, но ", 😺, " - это слишком мало.")
```

Output:
```
"Я не знаю, но 10 - это слишком мало."
```

Мы также можем использовать `*` для конкатенации!

```julia
s3*s4
```

Output:
```
"Сколько кошек, неужели их слишком много?"
```

<br>

---

<br>

## 4️⃣ Символы, байты и индексация строк

Это самый важный раздел занятия — и то место, где Julia ведёт себя **не так**,
как Python или JavaScript. Разберёмся спокойно.

### Строка — это последовательность байтов

Julia хранит строки в кодировке **UTF-8**. В ней символ занимает **разное**
число байтов: латинская буква — один, кириллическая — два, эмодзи — четыре.

Отсюда следует главное: у строки есть **две разные длины**.

```julia
en = "Hello"
ru = "Привет"

println("en: ", length(en), " символов, ", ncodeunits(en), " байт")
println("ru: ", length(ru), " символов, ", ncodeunits(ru), " байт")
```

Output:
```
en: 5 символов, 5 байт
ru: 6 символов, 12 байт
```

- `length(s)` — сколько в строке **символов**;
- `ncodeunits(s)` — сколько в строке **байтов** (единиц кодирования).

Для строки из латиницы эти числа совпадают. Для «Привет» — нет: шесть символов
занимают двенадцать байтов.

### Индекс в Julia — это номер байта, а не номер символа

Вот ключевое отличие от Python. В Python `s[1]` — «второй символ».
В Julia `s[i]` — «символ, начинающийся с байта номер `i`».

Пока строка состоит из латиницы, разницы не видно:

```julia
en = "Hello"

@show en[1]
@show en[2]
@show en[5]
```

Output:
```
en[1] = 'H'
en[2] = 'e'
en[5] = 'o'
'o': ASCII/Unicode U+006F (category Ll: Letter, lowercase)
```

А теперь то же самое с кириллицей. Первый символ достаётся нормально:

```julia
ru = "Привет"

@show ru[1]     # 'П' начинается с байта 1 — всё в порядке
```

Output:
```
ru[1] = 'П'
'П': Unicode U+041F (category Lu: Letter, uppercase)
```

Но байт номер 2 — это **вторая половина** буквы «П». Отдельным символом он не
является, и Julia откажется его возвращать:

```julia
ru[2]     # ← намеренная ошибка: байт 2 находится ВНУТРИ буквы «П»
```

Output:
```
@ Base ./strings/string.jl:464
 [4] top-level scope
   @ In[17]:1
```

Выше — **намеренная ошибка**. `StringIndexError` означает ровно одно: индекс
не указывает на начало символа.

Julia могла бы молча вернуть половину буквы или медленно пересчитать индекс в
номер символа. Она не делает ни того, ни другого: неверный индекс — это ошибка,
а индексация остаётся операцией за $O(1)$.

> ⚠️ **Отсюда следует практическое правило.**
> Такой цикл **неверен** для любого текста, кроме чисто латинского:
>
> ```julia
> for i in 1:length(s)     # ❌ не делайте так
>     println(s[i])
> end
> ```
>
> Для «Привет» он свалится на `i = 2`, потому что `length` считает символы,
> а индексация работает с байтами. Ниже — три правильных способа.

### Способ 1: перебор символов (используйте его по умолчанию)

Если индексы вам не нужны — а нужны они редко — просто перебирайте строку:

```julia
for c in ru
    print(c, " ")
end
println()

# Julia сама переходит от символа к символу, независимо от их длины в байтах
```

Output:
```
П р и в е т
```

### Способ 2: `eachindex`, когда индексы всё-таки нужны

`eachindex(s)` выдаёт **только корректные** индексы — те, что указывают на
начало символа:

```julia
@show collect(eachindex(ru))    # 1, 3, 5, 7, 9, 11 — по два байта на символ

for i in eachindex(ru)
    println("байт ", lpad(i, 2), " → ", ru[i])
end
```

Output:
```
collect(eachindex(ru)) = [1, 3, 5, 7, 9, 11]
байт  1 → П
байт  3 → р
байт  5 → и
байт  7 → в
байт  9 → е
байт 11 → т
```

Обратите внимание: индексы идут через один — 1, 3, 5, … — а вовсе не подряд.
Именно поэтому `1:length(s)` и не работает.

### Способ 3: ручная навигация — `firstindex`, `lastindex`, `nextind`, `prevind`

Иногда нужно двигаться по строке вручную. Для этого есть четыре функции:

```julia
@show firstindex(ru)    # индекс первого символа — всегда 1
@show lastindex(ru)     # индекс ПОСЛЕДНЕГО СИМВОЛА (11), а не длина строки!
@show ncodeunits(ru)    # байтов всего 12

# nextind / prevind перешагивают ровно через один символ
i = firstindex(ru)
@show ru[i]
i = nextind(ru, i);  @show ru[i]      # следующий символ
i = nextind(ru, i);  @show ru[i]
i = prevind(ru, i);  @show ru[i]      # и обратно
```

Output:
```
firstindex(ru) = 1
lastindex(ru) = 11
ncodeunits(ru) = 12
ru[i] = 'П'
ru[i] = 'р'
ru[i] = 'и'
ru[i] = 'р'
'р': Unicode U+0440 (category Ll: Letter, lowercase)
```

> 📌 **`lastindex(s)` — это не `length(s)`.** Первое — индекс последнего символа
> (для «Привет» это 11), второе — количество символов (6). Совпадают они только
> для строк из однобайтовых символов.
>
> Кстати, `s[end]` работает всегда: `end` внутри индексации — это как раз
> `lastindex(s)`.

```julia
@show ru[end]
@show ru[end] == ru[lastindex(ru)]
```

Output:
```
ru[end] = 'т'
ru[end] == ru[lastindex(ru)] = true
true
```

### `collect` — получить символы массивом

`collect(s)` превращает строку в `Vector{Char}`. Это удобно для разбора и
демонстраций, но помните: **создаётся новый массив**, то есть выделяется память.
Для простого прохода по строке цикл `for c in s` дешевле.

```julia
chars = collect(ru)

@show chars
@show typeof(chars)
@show length(chars) == length(ru)     # символов столько же
@show chars[2]                        # а вот здесь индекс — уже номер СИМВОЛА
```

Output:
```
chars = ['П', 'р', 'и', 'в', 'е', 'т']
typeof(chars) = Vector{Char}
length(chars) == length(ru) = true
chars[2] = 'р'
'р': Unicode U+0440 (category Ll: Letter, lowercase)
```

Сравните: `chars[2]` — это второй **символ** (`'р'`), а `ru[2]` — ошибка.
Разные структуры данных, разные правила индексации.

### Итог раздела

| Что нужно | Как делать | Как не делать |
|---|---|---|
| Пройти по символам | `for c in s` | `for i in 1:length(s)` |
| Пройти с индексами | `for i in eachindex(s)` | `for i in 1:ncodeunits(s)` |
| Число символов | `length(s)` | `ncodeunits(s)` |
| Число байтов | `ncodeunits(s)` | `length(s)` |
| Последний символ | `s[end]` | `s[length(s)]` |
| Соседний символ | `nextind` / `prevind` | `i + 1` |

> **Почему так, а не «как в Python».** Байтовая индексация даёт доступ за
> постоянное время. Если бы `s[i]` означало «i-й символ», Julia была бы вынуждена
> при каждом обращении просматривать строку с начала, и цикл по строке стал бы
> квадратичным. Julia выбрала скорость и явную ошибку вместо скрытой медленности.

<br>

---

<br>

## Упражнения 🚀

####  ✅  Задание 2.1 

Создайте строку, в которой "hi" повторяется 1000 раз, сначала с помощью "repeat", а затем с помощью оператора возведения в степень, который может вызывать "*". <br>
Присвойте ей переменную "hi", указанную ниже.

```julia
# Ваше решение:


```

**Правильное решение:**

<!-- grade_id: cell-5b38bdb1115161b4, points: 0, grade: false, solution: true -->
```julia
# Правильное решение

# Использование функции repeat
hi1 = repeat("hi", 1000)

# Использование оператора возведения в степень (^)
hi2 = "hi"^1000

# Проверка на равенство
@assert hi1 == hi2

# Присвоение переменной hi
hi = hi1  # или hi = hi2, так как они идентичны

```

Output:
```
"hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"[93m[1m ⋯ 1460 bytes ⋯ [22m[39m"hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"
```

<!-- grade_id: cell-20b1027a968fbb48, points: 1, grade: true, solution: false -->
```julia
@assert hi == "hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"
println(hi)
```

Output:
```
hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi
```

<br>

####  ✅ Задание 2.2 

Объявите две переменные

```julia
a = 3
b = 4
```
и используйте их для создания двух строк:
```julia
"3 + 4"
"7"
```
и сохраните результаты в `c` и `d` соответственно

```julia
# Ваше решение:


```

**Правильное решение:**

<!-- grade_id: cell-36beb524c8ecd33b, points: 0, grade: false, solution: true -->
```julia
# Правильное решение

# Объявление переменных
a = 3
b = 4

# Создание строк
c = "$a + $b"   # строка "3 + 4"
d = "$(a + b)"  # строка "7"

# Проверка корректности
@assert c == "3 + 4"
@assert d == "7"

println(c)
println(d)

```

Output:
```
3 + 4
7
```

<!-- grade_id: cell-14c3c5b015a97892, points: 1, grade: true, solution: false -->
```julia
@assert c == "3 + 4"
@assert d == "7"

println(c)
println(d)
```

Output:
```
3 + 4
7
```

<br>

#### ✅ Задание 2.3 — Символы и байты

Дана строка `phrase`, в которой намеренно смешаны кириллица, латиница и эмодзи.

1. В `n_chars` запишите количество **символов**.
2. В `n_bytes` запишите количество **байтов**.
3. В `idxs` запишите вектор всех **корректных индексов** строки.
4. В `last_char` запишите последний символ строки.
5. В `bad_index` запишите `true`, если обращение `phrase[2]` вызывает ошибку,
   и `false` в противном случае. Получите это значение **проверкой**, а не
   догадкой — обработайте исключение.

Подсказка к пункту 5: конструкция

```julia
try
    ...
catch e
    ...
end
```

позволяет перехватить исключение и проверить его тип с помощью `isa`.

```julia
phrase = "Привет, Julia! 🚀"
```

Output:
```
"Привет, Julia! 🚀"
```

```julia
# Ваше решение
```

**Правильное решение:**

```julia
# Правильное решение

n_chars = length(phrase)        # символы
n_bytes = ncodeunits(phrase)    # байты

# eachindex выдаёт только те индексы, которые указывают на начало символа
idxs = collect(eachindex(phrase))

last_char = phrase[end]         # end == lastindex(phrase)

# Проверяем реально, а не предполагаем: 'П' занимает два байта,
# поэтому индекс 2 попадает внутрь символа.
bad_index = try
    phrase[2]
    false                       # ошибки не было
catch e
    e isa StringIndexError      # true, если ошибка именно та, которую мы ждём
end

println("символов : ", n_chars)
println("байтов   : ", n_bytes)
println("индексов : ", length(idxs), " → ", idxs)
println("последний: ", last_char)
println("phrase[2] даёт StringIndexError: ", bad_index)
```

Output:
```
символов : 16
байтов   : 25
индексов : 16 → [1, 3, 5, 7, 9, 11, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22]
последний: 🚀
phrase[2] даёт StringIndexError: true
```

Обратите внимание: `length(idxs) == n_chars`, а `n_bytes` заметно больше.
Эмодзи `🚀` занимает целых четыре байта — это видно по «скачку» в конце
списка индексов.

```julia
@assert n_chars == 16                     "символов должно быть 16"
@assert n_bytes == 25                     "байтов должно быть 25"
@assert length(idxs) == n_chars           "корректных индексов ровно столько же, сколько символов"
@assert idxs == collect(eachindex(phrase))
@assert last_char == '🚀'
@assert bad_index === true                "phrase[2] должен приводить к StringIndexError"
@assert phrase[idxs[2]] == 'р'            "второй символ — строчная 'р'"
println("✔ Задание 2.3 выполнено")
```

Output:
```
✔ Задание 2.3 выполнено
```

---

[← Первое занятие. Знакомство с Julia](01%20-%20Getting%20started.ipynb) | [Оглавление](README.md) | [Третье занятие. Структуры данных →](03%20-%20Data%20structures.ipynb)