# **🚀 Структуры данных (Data structures). Урок Третий.**

Как только мы начнем работать с большим количеством фрагментов данных одновременно, нам будет удобно хранить данные в таких структурах, как массивы или словари (вместо того, чтобы полагаться только на переменные).

## **📌 Темы:**

1. **Кортежи (Tuples).**
   
2. **Словари (Dictionaries).**
   
3. **Массивы (Arrays).**

<br>

В общем, и кортежи, и массивы - это упорядоченные последовательности элементов (поэтому мы можем индексировать их). Словари и массивы могут изменяться.
Ниже, мы рассмотрим это подробнее !

<br>

---

<br>


## 🎯 После этого занятия вы сможете

- различать `Tuple`, `NamedTuple`, `Dict`, `Vector` и `Matrix`;
- объяснять, какие из них изменяемы, а какие нет;
- выбирать структуру данных под задачу;
- понимать, чем `Vector{Any}` отличается от `Vector{Int}` по производительности,
  и когда разнородность оправданна.

### **1️⃣ Кортежи (Tuples)**


Мы можем создать **кортеж**, заключив упорядоченную коллекцию элементов в `( )`.

**Синтаксис:** <br>
```julia
(item1, item2, ...)
```


```julia
myfavoriteanimals = ("penguins", "cats", "sugargliders") # Это кортеж (tuple)
```

Output:
```
("penguins", "cats", "sugargliders")
```

Мы можем проиндексировать этот **кортеж**,

```julia
myfavoriteanimals[1]
```

Output:
```
"penguins"
```

но поскольку **кортежи** неизменяемы, мы не можем их обновить !

```julia
myfavoriteanimals[1] = "otters"
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[3]:1
```

<br>

### **🔹 NamedTuples - "Именованные кортежи".**

В Julia **NamedTuple** — это кортеж, где каждому элементу присваивается имя, такие кортежи принято называть - **"Именованные кортежи"**.

Они доказывают свою полезность, когда нужно хранить структурированные данные с удобным доступом по именам.

У них особый синтаксис, использующий `=` внутри кортежа:


` (name1 = item1, name2 = item2, ...) `

<br>

Пример:

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Output:
```
(bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

📌 Как и обычные  кортежи `Tuples`, именованные кортежи `NamedTuples` упорядочены, поэтому мы можем извлекать их элементы с помощью индексации:

```julia
myfavoriteanimals[1]
```

Output:
```
"penguins"
```

📌 Они также добавляют специальную возможность доступа к значениям по их названию:

```julia
myfavoriteanimals.bird
```

Output:
```
"penguins"
```

**Пример синтаксиса NamedTuple:**

```julia
# Создание NamedTuple
person = (name = "Alice", age = 30, city = "New York")

# Доступ к элементам по имени
println(person.name)  # Выведет: Alice
println(person.age)   # Выведет: 30
println(person.city)  # Выведет: New York

# Можно обращаться как к обычному кортежу
println(person[1])  # Выведет: Alice

# Деструктуризация NamedTuple
(name, age, city) = person
println(name)  # Alice
println(age)   # 30
println(city)  # New York

```

Output:
```
Alice
30
New York
Alice
Alice
30
New York
```

📌 **Создание NamedTuple с переменными:**

```julia
name = "Bob"
age = 25
city = "Los Angeles"

person2 = (name = name, age = age, city = city)  # Автоматическая подстановка значений

println(person2)  # (name = "Bob", age = 25, city = "Los Angeles")

```

Output:
```
(name = "Bob", age = 25, city = "Los Angeles")
```

📌 **Изменение NamedTuple**

NamedTuple — неизменяемая структура данных, но можно создать новую версию с изменёнными значениями:

```julia
person_updated = merge(person, (age = 31,))  # Изменение возраста
println(person_updated)  # (name = "Alice", age = 31, city = "New York")

```

Output:
```
(name = "Alice", age = 31, city = "New York")
```

📌 **Применение в функциях:**

```julia
function greet(person)
    println("Hello, $(person.name)! You are $(person.age) years old and live in $(person.city).")
end

greet(person)  # Выведет: Hello, Alice! You are 30 years old and live in New York.

```

Output:
```
Hello, Alice! You are 30 years old and live in New York.
```

💡 **NamedTuple**  удобны, когда нужна структура данных без избыточности, но с возможностью удобного доступа по именам.

---
<br>

### **2️⃣ Словари  (Dictionaries).**

Если у нас есть наборы данных, связанные друг с другом, мы можем сохранить эти данные в словаре. Мы можем создать словарь, используя функцию `Dict()`, которую мы можем инициализировать как пустой словарь или как один, хранящий пары ключ-значение.

**Синтаксис:**

```julia
Dict(key1 => value1, key2 => value2, ...)
```

Хорошим примером является список контактов, где мы связываем имена с телефонными номерами, как показано ниже:


```julia
myphonebook = Dict("Света" => "867-5309", "Гена" => "555-2368")
```

Output:
```
Dict{String, String} with 2 entries:
  "Света" => "867-5309"
  "Гена"  => "555-2368"
```

В этом примере каждое имя и номер — это пара «ключ» и «значение». 

📌 Мы можем получить **номер** Светы (т.е. **значение**), используя связанный ключ:

```julia
myphonebook["Света"]
```

Output:
```
"867-5309"
```

📌 Мы можем **добавить** еще одну запись в этот словарь следующим образом:

```julia
myphonebook["Сергей"] = "09-565-997"
```

Output:
```
"09-565-997"
```

Давайте проверим, как выглядит наша телефонная книга сейчас...

```julia
myphonebook
```

Output:
```
Dict{String, String} with 3 entries:
  "Света"  => "867-5309"
  "Гена"   => "555-2368"
  "Сергей" => "09-565-997"
```

📌 Мы можем **удалить** Сергей из нашего списка контактов и одновременно получить его номер, используя `pop!` 

```julia
pop!(myphonebook, "Сергей")
```

Output:
```
"09-565-997"
```

Проверим нашу записную книжку:

```julia
myphonebook
```

Output:
```
Dict{String, String} with 2 entries:
  "Света" => "867-5309"
  "Гена"  => "555-2368"
```

💡 В отличие от кортежей и массивов, ***словари не упорядочены.*** 

- ***Поэтому мы не можем индексировать их !!!***

```julia
myphonebook[1] # Ошибка
```

Output:
```
@ Base ./dict.jl:477
 [2] top-level scope
   @ In[17]:1
```

В приведенном выше примере `julia` думает, что вы пытаетесь получить доступ к значению, связанному с ключом `1`.

<br>

---

<br>

### **3️⃣ Массивы (Arrays).**

В отличие от кортежей, массивы изменяемы. В отличие от словарей, массивы содержат упорядоченные коллекции. <br>
Мы можем создать массив, заключив эту коллекцию в `[ ]`.

**Синтаксис:** <br>

` [item1, item2, ...] `

<br>

**Рассмотрим на примерах:**

#### 📌 Создадим одномерный простой массив c именами друзей :

```julia
# одномерный массив

myfriends_men = ["Толя", "Роберт", "Борис", "Геннадий", "Михаил"]


```

Output:
```
5-element Vector{String}:
 "Толя"
 "Роберт"
 "Борис"
 "Геннадий"
 "Михаил"
```

💡 В Julia `5-element Vector{String}` означает, что переменная `myfriends_men` является **вектором** (одномерным массивом) длины **5**, состоящим из элементов типа `String` - из строк.

<br>

#### 📌 Теперь создадим одномерный **вложенный массив**, по сути "список списков":

```julia
# вложенный массив

myfriends = [["Толя", "Роберт", "Борис", "Геннадий", "Михаил"],["Антонина", "Роксана", "Барбара", "Лилия", "Мария"]]
```

Output:
```
2-element Vector{Vector{String}}:
 ["Толя", "Роберт", "Борис", "Геннадий", "Михаил"]
 ["Антонина", "Роксана", "Барбара", "Лилия", "Мария"]
```

💡  `5` в `{String,1}` означает, что это одномерный вектор. `Array{String,2}` будет двумерной матрицей и т. д. `String` — это тип каждого элемента.

или массив для хранения последовательности чисел Фибоначчи:

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13] # Массив чисел Фибоначчи
```

Output:
```
7-element Vector{Int64}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
print(fibonacci) # Выведет: [1, 1, 2, 3, 5, 8, 13]
```

Output:
```
[1, 1, 2, 3, 5, 8, 13]
```


#### 📌 Как создать **двумерный массив - Matrix** в Julia?

Массивы в **Julia** могут быть как **одномерными**, так и **многомерными - Matrix**, рассмотрим это на примере двумерного массива. 

В Julia двумерный массив - **Matrix**, создаётся следующим образом:

```julia
myfriends_matrix = [
    "Толя"      "Роберт"    "Борис"    "Геннадий"    "Михаил";
    "Антонина"  "Роксана"   "Барбара"  "Лилия"       "Мария"
]

```

Output:
```
2×5 Matrix{String}:
 "Толя"      "Роберт"   "Борис"    "Геннадий"  "Михаил"
 "Антонина"  "Роксана"  "Барбара"  "Лилия"     "Мария"
```

🔹 Теперь `myfriends_matrix` представляет собой настоящий `Matrix{String}` с размерностью 2х5 - две строки и пять столбцов!

<br>

📌 **Разница между вложенным одномерным массивом и двумерным массивом Matrix :**

`Vector{Vector{String}}` → список списков (разная длина возможна, не обязательно прямоугольная структура).

`Matrix{String}` → настоящий двумерный массив c фиксированным размером (в нашем случае 2х5).

<br>

#### 📌 **Смешанные (гетерогенные) массивы в Julia.**

Как мы уже рассмотрели ранее В Julia массив **обычно создаётся с однородным типом**, например:  

```julia
arr = [1, 2, 3, 4]  # Это Vector{Int64}, потому что все элементы - числа
```

Output:
```
4-element Vector{Int64}:
 1
 2
 3
 4
```

Но если нам необходимо создать массив, в котором будут **разные типы данных**?

В этом случае Julia **автоматически создаёт массив типа `Vector{Any}`**, который может содержать **любые типы элементов**, такие маасивы называются **ГЕТЕРОГЕННЫМИ**.

**Пример:**

```julia
mixture = [1, 1, 2, 3, "Толя", "Роберт"]  # Это Vector{Any}, потому что элементы разные
```

Output:
```
6-element Vector{Any}:
 1
 1
 2
 3
  "Толя"
  "Роберт"
```

Проверяем тип:

```julia
typeof(mixture)  # Выведет: Vector{Any}
```

Output:
```
Vector{Any}[90m (alias for [39m[90mArray{Any, 1}[39m[90m)[39m
```

**🔹 Разбираем выражение**
- **Числа `1, 1, 2, 3`** → имеют тип `Int64` (или `Int32` в 32-битной системе).
- **Строки `"Ted", "Robyn"`** → имеют тип `String`.
- **Julia автоматически выбирает тип `Vector{Any}`**, так как типы данных **смешаны**.

**Вывод структуры массива:**

```julia
eltype(mixture)  # Any
```

Output:
```
Any
```

```julia
length(mixture)  # 6
```

Output:
```
6
```

- **`eltype(mixture) == Any`** → массив содержит **разные типы данных**.
- **`length(mixture) == 6`** → в массиве **6 элементов**.

##### **🔹 Как Julia решает, какой тип использовать в массиве?**

**1️⃣ Если все элементы одного типа**, Julia создаёт **однородный массив**:

```julia
nums = [1, 2, 3, 4]  
typeof(nums)  # Vector{Int64}
```

Output:
```
Vector{Int64}[90m (alias for [39m[90mArray{Int64, 1}[39m[90m)[39m
```

**2️⃣ Если элементы имеют разные типы**, Julia **автоматически создаёт `Vector{Any}`**:

```julia
mixed = [1, "Hello", 3.14, true]
typeof(mixed)  # Vector{Any}
```

Output:
```
Vector{Any}[90m (alias for [39m[90mArray{Any, 1}[39m[90m)[39m
```


**3️⃣ Если типы совместимы, Julia выбирает наибольший общий тип**:

```julia
mixed_numbers = [1, 2.5, 3]  
typeof(mixed_numbers)  # Vector{Float64}, так как Int можно привести к Float
```

Output:
```
Vector{Float64}[90m (alias for [39m[90mArray{Float64, 1}[39m[90m)[39m
```

##### **⚖️ `Vector{Any}`: чем приходится платить**

`Vector{Any}` — **не ошибка и не признак плохого кода**. Это осознанный
компромисс, и полезно понимать обе его стороны.

**Что вы теряете.** Когда тип элемента — `Any`, компилятор не знает заранее,
что лежит в ячейке. Из этого следует:

- **упаковка (boxing)** — в массиве хранятся не сами значения, а ссылки на них;
- **динамическая диспетчеризация** — тип каждого элемента выясняется во время
  выполнения (подробнее — в занятии 9);
- **дополнительные выделения памяти** и, как следствие, нагрузка на сборщик мусора;
- **потеря специализации** — компилятор не может сгенерировать узкий машинный код.

**Когда это действительно важно.** Только если массив находится на
**горячем пути** — в числовом цикле, который выполняется миллионы раз.
Разницу мы измерим ниже.

**Когда это неважно.** Если в массиве десяток разнородных значений, которые вы
обходите один раз, — накладные расходы незаметны, а `Any` честно отражает
структуру данных.

Не будем верить на слово — измерим. Один и тот же расчёт на однородном и на
разнородном массиве:

```julia
using BenchmarkTools

n = 100_000
typed   = collect(1:n)              # Vector{Int64}
untyped = Any[i for i in 1:n]       # Vector{Any} с теми же числами

# Функция одна и та же; отличается только тип аргумента
function total(v)
    s = 0
    for x in v
        s += x
    end
    s
end

@show eltype(typed), eltype(untyped)
@show total(typed) == total(untyped)     # результат одинаковый

@btime total($typed)
@btime total($untyped)
```

Output:
```
(eltype(typed), eltype(untyped)) = (Int64, Any)
total(typed) == total(untyped) = true
  40.515 μs (0 allocations: 0 bytes)
  3.058 ms (99969 allocations: 1.53 MiB)
5000050000
```

Результат один и тот же, время — нет. Обратите внимание и на колонку
`allocations`: у варианта с `Any` их много, у типизированного — нет.

Но заметьте: чтобы разница стала видна, понадобилось **сто тысяч** элементов.
На массиве из десяти значений вы её не измерите.

##### **🎯 Золотая середина: небольшой `Union`**

Между «все элементы одного конкретного типа» и «что угодно» есть важный
промежуточный вариант — **небольшое объединение типов**.

Julia обрабатывает такие типы особым образом. Если в `Union` входит несколько
конкретных типов, компилятор применяет приём под названием
*union splitting*: он генерирует ветку под каждый вариант и выбирает нужную
дешёвой проверкой. Это **намного** быстрее, чем полностью динамический `Any`.

Самый частый случай в реальных данных — пропущенные значения:

```julia
# Измерения, в которых часть значений отсутствует.
# Тип элемента — Union{Missing, Float64}, а НЕ Any.
measurements = [1.5, 2.0, missing, 3.25, missing, 4.0]

@show eltype(measurements)
@show typeof(measurements)
```

Output:
```
eltype(measurements) = Union{Missing, Float64}
typeof(measurements) = Vector{Union{Missing, Float64}}
Vector{Union{Missing, Float64}}[90m (alias for [39m[90mArray{Union{Missing, Float64}, 1}[39m[90m)[39m
```

Julia вывела тип `Union{Missing, Float64}` — она не «сдалась» до `Any`.
Это ровно то, что нужно: два конкретных варианта, и оба известны компилятору.

Работать с такими данными помогают `skipmissing`, `ismissing` и `coalesce`:

```julia
@show count(ismissing, measurements)          # сколько пропусков
@show collect(skipmissing(measurements))      # только известные значения
@show sum(skipmissing(measurements))          # сумма без пропусков
@show coalesce.(measurements, 0.0)            # заменить пропуски нулями
```

Output:
```
count(ismissing, measurements) = 2
collect(skipmissing(measurements)) = [1.5, 2.0, 3.25, 4.0]
sum(skipmissing(measurements)) = 10.75
coalesce.(measurements, 0.0) = [1.5, 2.0, 0.0, 3.25, 0.0, 4.0]
6-element Vector{Float64}:
 1.5
 2.0
 0.0
 3.25
 0.0
 4.0
```

Обратите внимание на поведение `missing` в арифметике — оно «заразно» по замыслу:

```julia
@show 1 + missing            # missing: результат неизвестен
@show missing == missing     # тоже missing, а не true!
@show ismissing(missing)     # вот так проверяют на пропуск
```

Output:
```
1 + missing = missing
missing == missing = missing
ismissing(missing) = true
true
```

> 📌 **Почему `missing == missing` даёт `missing`, а не `true`.**
> `missing` означает «значение неизвестно». Два неизвестных значения могут
> оказаться как равными, так и разными — поэтому честный ответ «неизвестно».
> Для проверки самого факта пропуска есть `ismissing`.

Сравним три подхода к одним и тем же данным:

| Тип элемента | Когда уместен | Цена |
|---|---|---|
| `Vector{Float64}` | пропусков нет | нулевая, самый быстрый вариант |
| `Vector{Union{Missing, Float64}}` | пропуски возможны | небольшая, union splitting |
| `Vector{Any}` | типы заранее неизвестны | заметная на горячем пути |

**Явное указание типа** — способ сказать, что вы имеете в виду:

```julia
# Разрешаем ровно два типа — не больше
limited = Union{Int, String}[1, 2, "Толя", "Роберт"]
@show eltype(limited)

# Явный Any — когда разнородность действительно нужна
freeform = Any[1, 2, "Фаина", 3.14, true]
@show eltype(freeform)
```

Output:
```
eltype(limited) = Union{Int64, String}
eltype(freeform) = Any
Any
```

##### **😺 Выводы:**

- **`Vector{Any}` возникает автоматически**, если элементы разных типов.

- **Однородные массивы (`Vector{Int64}`, `Vector{Float64}`) быстрее** — компилятор
  знает тип и может специализировать код.

- **Это не значит, что `Any` — плохо.** Если данные действительно разнородны,
  `Any` их честно описывает. Проблема возникает только на горячем пути.

- **Между «один конкретный тип» и `Any` есть середина** — небольшой `Union`,
  например `Union{Int, String}` или `Union{Missing, Float64}`. Julia умеет
  работать с такими типами эффективно; следующий раздел — как раз об этом.

<br>

> 📌 **Порядок действий, а не догма.**
> Сначала выберите структуру данных, которая **правильно описывает вашу
> предметную область**. Оптимизируйте представление тогда, когда измерение
> покажет, что этот участок кода действительно на горячем пути.

<br>

---

<br>

### **🔹 Доступ к элементам массива в Julia**  

В Julia доступ к элементам массива осуществляется с **индексацией, срезами (slicing) и фильтрацией**.  

<br>

📌 **Важно помнить!** 

В отличие от Python, где индексация начинается с `0`, в **Julia индексация начинается с `1`**!  

<br>


#### **1️⃣ Индексация элементов**  
**Синтаксис:** `array[index]`  

```julia
arr = ["Толя", "Роберт", "Борис", "Геннадий", "Михаил"]     # Массив строк

println(arr[1])  # "Толя" (первый элемент)
println(arr[5])  # "Михаил" (пятый элемент)
```

Output:
```
Толя
Михаил
```

📌 **`arr[end]` возвращает последний элемент:**

```julia
println(arr[end])  # "Михаил" (последний элемент)
```

Output:
```
Михаил
```


#### **2️⃣ Доступ к элементам вложенного массива (`Vector{Vector{T}}`)**
Если у нас **список списков** (`Vector{Vector{String}}`):

```julia
nested_arr = [["Толя", "Роберт"], ["Антонина", "Роксана"]]
println(nested_arr[1])  # ["Толя", "Роберт"]
println(nested_arr[2][1])  # "Антонина" (1-й элемент 2-го подсписка)
```

Output:
```
["Толя", "Роберт"]
Антонина
```

📌 **Если бы это была матрица (`Matrix{String}`), доступ был бы через `arr[row, col]` (arr[строка,столбец]).**

#### **3️⃣ Срезы (slicing)**

Срез (или подмассив) создаётся через `start:stop`:

```julia
println(arr[2:4])  # ["Роберт", "Борис", "Геннадий"] (2-й, 3-й и 4-й элементы)
```

Output:
```
["Роберт", "Борис", "Геннадий"]
```

📌 **Диапазон `2:4` включает оба конца!**  
<br>

Другие варианты:

```julia
println(arr[1:end-1])  # Все элементы, кроме последнего
println(arr[1:2:end])  # Каждый второй элемент (шаг 2)
```

Output:
```
["Толя", "Роберт", "Борис", "Геннадий"]
["Толя", "Борис", "Михаил"]
```

#### **4️⃣ Выбор элементов с условием (фильтрация)**

Фильтрация элементов с `filter`:

```julia
println(filter(x -> length(x) > 5, arr))  # ["Роберт", "Геннадий", "Михаил"]
```

Output:
```
["Роберт", "Геннадий", "Михаил"]
```

📌 **Оставляем только имена длиной больше 5 символов.**

#### **5️⃣ Изменение элементов массива**

```julia
arr[1] = "Иван"  # Теперь первый элемент "Иван"
println(arr)
```

Output:
```
["Иван", "Роберт", "Борис", "Геннадий", "Михаил"]
```

📌 **В Julia можно изменять элементы массива, если его тип позволяет, т.е. для `Vector{Any}` или `Vector{String}`.**

<br>

##### **😺 Выводы:**

- Julia использует индексацию с `1`, а не с `0`. 

- Доступ к вложенным массивам требует `[i][j]`.

- Срезы работают через `start:stop`, включая оба конца.  

- Можно фильтровать элементы с `filter` и изменять их напрямую.  



<br>

---

<br>

### **🔹 Преобразование в массив определенного типа**  

Иногда в Julia необходимо преобразовать **гетерогенный массив (`Vector{Any}`)** в массив с **определённым типом** (`Vector{Int}`, `Vector{String}` и т. д.).  

Если все числа должны быть **целыми (`Int`)**, а строки хранятся отдельно, можно создать два отдельных массива:

```julia
mixed_array = [1, 2, "Привет", 3.5, "Julia"]
int_array = [x for x in mixed_array if x isa Int]  # Отбираем только Int
string_array = [x for x in mixed_array if x isa String]  # Отбираем только String

println(int_array)   # [1, 2]
println(string_array) # ["Привет", "Julia"]
```

Output:
```
[1, 2]
["Привет", "Julia"]
```


📌 **Это оптимальнее, чем `Vector{Any}`**, так как Julia эффективнее работает с **однородными массивами** (`Vector{Int}` или `Vector{String}`).  

#### **Лучшие альтернативы `Vector{Any}`**

Вместо хранения разнородных данных в `Vector{Any}`, можно использовать:

- **`NamedTuple`** (кортеж с именованными полями)
  
- **`Dict{Symbol, Any}`** (словарь)
  
- **`Struct`** (структура данных с определёнными типами)
  
  <br>

1. Пример с **`NamedTuple`**:

```julia
person = (name = "Антон", age = 30, height = 1.7)
println(person.name)  # "Антон"
println(person.age)   # 30
```

Output:
```
Антон
30
```

2. Пример с **`Dict`**:

```julia
data = Dict(:name => "Рома", :age => 25, :score => 95.7)
println(data[:name])  # "Рома"
```

Output:
```
Рома
```

3. Пример со **`Struct`**:

```julia
struct Person
    name::String
    age::Int
    height::Float64
end

p = Person("Карл", 28, 1.75)
println(p.name)  # "Карл"
```

Output:
```
Карл
```

📌 **Если элементы разнородны, но их состав фиксирован и осмыслен, `NamedTuple`
или `struct` описывают данные точнее, чем `Vector{Any}`:** у полей появляются
имена и типы, а компилятор снова видит конкретную структуру.

`Vector{Any}` уместен там, где набор элементов заранее неизвестен и однороден
только по назначению, а не по типу.

<br>

---

<br>

### **🔹 Многомерные массивы в Julia**

До сих пор мы рассматривали **одномерные массивы** и только слегка затронули **многомерные массивы**, но Julia поддерживает **массивы произвольной размерности**.

<br>

#### **Примеры многомерных массивов**


```julia
rand(4, 3)  # Двумерный массив 4x3 (случайные числа)

```

Output:
```
4×3 Matrix{Float64}:
 0.140012  0.369142  0.55043
 0.99772   0.992709  0.359451
 0.384993  0.196042  0.785839
 0.254046  0.805817  0.102872
```

```julia
rand(4, 3, 2)  # Трёхмерный массив 4x3x2 (случайные числа)
```

Output:
```
4×3×2 Array{Float64, 3}:
[:, :, 1] =
 0.623206  0.410443   0.614653
 0.559508  0.85121    0.853861
 0.523496  0.0487672  0.127908
 0.927745  0.0174999  0.0146255

[:, :, 2] =
 0.617739  0.237805  0.643915
 0.640019  0.552543  0.485879
 0.1425    0.864079  0.486725
 0.462702  0.293697  0.842796
```

📌 **Julia поддерживает `N`-мерные массивы (`Array{T, N}`), где `N` — количество измерений.**  

**Пример 3D-массива (`Array{Float64, 3}`):**

```julia
A = rand(2, 2, 2)  # 3D-массив 2x2x2
println(size(A))    # (2, 2, 2)
```

Output:
```
(2, 2, 2)
```

- Первый индекс `i` — строка  
  
- Второй `j` — столбец  
  
- Третий `k` — глубина  

📌 **В Julia индексация многомерных массивов — `[i, j, k]`, а не `[i][j][k]`, как в Python.**  

<br>

#### **🔹 Осторожно с копированием массивов!**

В Julia **переменные — это ссылки на объекты в памяти**, а не сами данные.  

Поэтому, если вы **присваиваете один массив другому**, то **оба указывают на одни и те же данные!**  

##### **Пример: изменение массива через ссылку**

```julia
fibonacci = [1, 1, 2, 3, 5, 8]
mynumbers = fibonacci  # Не копия! Просто ссылка на тот же массив

mynumbers[1] = 404  # Изменяем somenumbers
println(fibonacci)  # [404, 1, 2, 3, 5, 8] - оригинальный массив тоже изменился!
```

Output:
```
[404, 1, 2, 3, 5, 8]
```

📌 **Это не копия, а просто другая ссылка на тот же массив!**  

<br>

#### **🔹 Как сделать настоящую копию массива?**

Используйте **`copy()`**, чтобы создать новый массив:

```julia
# Восстанавливаем массив Фибоначчи
fibonacci[1] = 1

# Создаём полную копию
mymorenumbers = copy(fibonacci)

# Изменяем копию
mymorenumbers[1] = 404

# Оригинальный массив не изменился!
println(fibonacci)  # [1, 1, 2, 3, 5, 8]
```

Output:
```
[1, 1, 2, 3, 5, 8]
```

📌 **Теперь `somemorenumbers` и `fibonacci` — два разных массива, и изменения в одном не затрагивают другой.**  

<br>

##### **😺 Выводы:**

- **`Vector{Any}` создаётся при смешанных типах.** Это не ошибка — это следствие
  того, что вы попросили.

- **Однородные массивы (`Vector{Int}`, `Vector{String}`) быстрее**, потому что
  компилятор может специализировать код.

- **Небольшой `Union`** (например, `Union{Missing, Float64}`) — эффективная
  середина между конкретным типом и `Any`.

- **Для фиксированного набора разнородных полей** лучше подходят `NamedTuple`,
  `Dict` или `struct` — там у данных появляются имена.

- **Julia поддерживает `N`-мерные массивы (`Array{T, N}`).**

- **Присваивание массива создаёт ссылку, а `copy()` — независимую копию.**

<br>

**Главное правило:** структура данных выбирается под **модель данных**.
Производительность — второй вопрос, и решается он измерением, а не суевериями.

### Упражнения 🚀

#### ✅ Задание 3.1

Создайте массив `a_ray` с помощью следующего кода:

```julia
a_ray = [1, 2, 3]
```

Добавьте число `4` в конец этого массива, а затем удалите его.

```julia
# Ваше решение


```

```julia
# Правильное решение:

# Создание массива
a_ray = [1, 2, 3]

# Добавление числа 4 в конец массива
push!(a_ray, 4)

# Удаление последнего элемента массива
pop!(a_ray)

```

Output:
```
4
```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3] # Проверка результата
```

<br>

#### ✅ Задание 3.2

Создайте **новый словарь** под названием "my_phonebook", в котором номер Евгения сохранен в виде целого числа 8675309, а номер  Константина - в виде строки "555-2368"


```julia
# Ваше решение


```

```julia
# Правильное решение:

# Номер Евгения — целое число, номер Константина — строка, поэтому
# словарь должен допускать значения разных типов: Dict{String, Any}
my_phonebook = Dict{String, Any}(
    "Евгений" => 8675309,
    "Константин" => "555-2368"
)

```

Output:
```
Dict{String, Any} with 2 entries:
  "Константин" => "555-2368"
  "Евгений"    => 8675309
```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert my_phonebook == Dict("Евгений" => 8675309, "Константин" => "555-2368") # Проверка результата
```

```julia
@assert typeof(my_phonebook) == Dict{String, Any} # Проверка типа
```

<br>

#### ✅ Задание 3.3


Добавьте ключ `Emergency` со значением `911` (целое число) в `flexible_phonebook`.

```julia
# Ваше решение


```

✅ Правильное решение:

<br>
Если изначально flexible_phonebook был создан как Dict{String, String} или Dict{String, Int}, то добавление значения другого типа вызовет ошибку. Чтобы этого избежать, нужно изначально создать словарь с типом Dict{String, Any}, чтобы он мог хранить как строки, так и числа.



```julia
# Создание словаря, который может хранить значения разных типов
flexible_phonebook = Dict{String, Any}(
    "Евгений" => 8675309,
    "Константин" => "555-2368"
)

# Добавление ключа "Emergency" со значением 911
flexible_phonebook["Emergency"] = 911


# Проверка результата
println(flexible_phonebook)  # Dict{String, Any}("Emergency" => 911, "Евгений" => 8675309, "Константин" => "555-2368")
```

Output:
```
Dict{String, Any}("Emergency" => 911, "Константин" => "555-2368", "Евгений" => 8675309)
```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
@assert haskey(flexible_phonebook, "Emergency")
```

<!-- grade_id: cell-c8c40ababe7cdfce, points: 2, grade: true, solution: false -->
```julia
@assert flexible_phonebook["Emergency"] == 911
```

<br>

### 🔥 Oператор **`@assert`** 

В Julia оператор **`@assert`** используется для проверки утверждений (**assertions**).  

##### 🔍 **Разбор кода:**


```julia
@assert flexible_phonebook["Emergency"] == 911
```

**Что делает этот код?**

- Он проверяет, что значение, хранящееся в словаре `flexible_phonebook` по ключу `"Emergency"`, **равно** `911`.
  
- Если выражение `flexible_phonebook["Emergency"] == 911` **истинно**, программа **продолжает выполнение** без ошибок.
  
- Если выражение **ложно**, Julia выдаст **ошибку AssertionError**, что указывает на проблему в коде.
  
<br>

##### ✅ **Пример с корректными данными (не вызовет ошибку):**
```julia
# Создание словаря
flexible_phonebook = Dict{String, Any}("Emergency" => 911)

# Проверка утверждения
@assert flexible_phonebook["Emergency"] == 911  # Всё верно, ошибок нет

println("Проверка пройдена!")  # Этот код выполнится, так как ошибки нет
```

<br>

##### ❌ **Пример с ошибкой:**
```julia
# Создание словаря с неверным значением
flexible_phonebook = Dict{String, Any}("Emergency" => 112)

# Проверка утверждения
@assert flexible_phonebook["Emergency"] == 911  # Ошибка! Значение 112 ≠ 911
```
**Вывод ошибки:**
```
ERROR: AssertionError: flexible_phonebook["Emergency"] == 911
```
Это означает, что в словаре `"Emergency"` хранит **не `911`**, а другое значение.

##### 🔹 **Вывод:**
- `@assert` помогает **отлавливать ошибки в коде** на ранних стадиях.
- Если условие истинно ✅ → программа продолжает работу.
- Если условие ложно ❌ → выбрасывается `AssertionError`, и выполнение останавливается.

Этот механизм полезен для **отладки, тестирования и контроля корректности данных**. 🚀

<br>

#### ✅ Задание 3.4

Почему мы можем добавить целое число в качестве значения в "flexible_phonebook", но не в "myphonebook"? Как мы могли инициализировать "мою телефонную книгу", чтобы она принимала целые числа в качестве значений? (подсказка: попробуйте воспользоваться [документацией Джулии по словарям](https://docs.julialang.org/en/v1/base/collections/#Dictionaries))

##### 🔍 **Почему в `flexible_phonebook` можно добавить `Int`, а в `myphonebook` — нет?**

В Julia **каждый словарь (`Dict`) имеет чётко определённые типы для ключей и значений**.  
Когда мы создаём `myphonebook` **без явного указания типа**, Julia автоматически **выводит тип данных из переданных значений**.

###### ⚠️ **Пример ошибки**:

```julia
# Создание словаря с строковыми значениями
myphonebook = Dict("Анна" => "1234", "Борис" => "5678")

# Попытка добавить число вызовет ошибку
myphonebook["Emergency"] = 911  # ❌ Ошибка!
```

Output:
```
@ Base ./dict.jl:355
 [2] top-level scope
   @ In[63]:5
```

##### ❌ **Почему возникает ошибка?**
- При создании `myphonebook` Julia автоматически определяет его тип как:

```julia
  Dict{String, String}  # Ключи — строки, значения — только строки
```

Output:
```
Dict{String, String}
```

- `911` — это `Int`, а словарь **не принимает значения другого типа**.
- В Julia **нельзя менять типы элементов после создания словаря**.

✅ **Как создать `myphonebook`, чтобы он принимал `Int` и `String`?**
Нужно **явно указать тип значений как `Any`** при создании:

```julia
# Разрешаем хранить любые типы значений (String, Int, и т. д.)
myphonebook = Dict{String, Any}("Анна" => "1234", "Борис" => "5678")

# Теперь можно добавить целое число
myphonebook["Emergency"] = 911  # ✅ Работает!

Dict{String, Any}  # Ключи — строки, значения — любые типы
```

Output:
```
Dict{String, Any}
```

Теперь словарь принимает разные типы значений:

```julia
println(myphonebook)  
# Вывод: Dict{String, Any}("Emergency" => 911, "Борис" => "5678", "Анна" => "1234")
```

Output:
```
Dict{String, Any}("Emergency" => 911, "Борис" => "5678", "Анна" => "1234")
```


##### 💡 **Выводы**:

- **Обычные словари (`Dict{String, String}`)** строго ограничены типами значений.
  
- **`Dict{String, Any}`** позволяет хранить **разные типы значений** (например, `Int` и `String`).
  
- Чтобы **избежать ошибок**, **явно указывайте `Dict{String, Any}`**, если вам нужно хранить разные типы данных.

🚀 Теперь `myphonebook` может хранить **и строки, и числа**!

---

[← Второе занятие. Строки](02%20-%20Strings.ipynb) | [Оглавление](README.md) | [Четвёртое занятие. Циклы →](04%20-%20Loops.ipynb)