# **🚀 Функции в Julia. Урок Шестой.**

<br>

## **📌 Темы:**

1. **Как объявить функцию.**
   
2. **Утиная типизация (Duck-typing) в Julia.**
   
3. **Изменяющие (Mutating) и неизменяющие (Non-mutating) функции.**
   
4. **Некоторые функции высшего порядка.**

<br>

---

 <br>

### **1️⃣ Как объявить функцию**

В Julia можно объявлять функции несколькими способами:

#### **1.1. Использование ключевого слова `function`**

```julia
function add(x, y)
    return x + y
end
```

Output:
```
add (generic function with 1 method)
```

**Вызываем функцию:**

```julia
add(3, 5)  # Вызываем функцию add с аргументами 3 и 5

# Вернет 8
```

Output:
```
8
```



#### **1.2. Однострочное объявление (синтаксис "assignment")**

Здесь `return` можно опустить, так как в Julia результатом функции становится последняя вычисленная строка:

```julia
add(x, y) = x + y # Тоже самое, что и выше
```

Output:
```
add (generic function with 1 method)
```

```julia
add(3, 5) # Вернет 8
```

Output:
```
8
```



#### **1.3. Анонимные функции (Lambda-выражения)**

```julia
square = x -> x^2 # Создаем анонимную функцию, которая возводит число в квадрате
square(4)  # Вернет 16
```

Output:
```
16
```

<br>

---

<br>

### **2️⃣ Утиная типизация (Duck-typing) в Julia**  
<br>

#### 📌 Что такое утиная типизация?

Утиная типизация (Duck-typing) — это концепция в языках программирования, в которых тип объекта определяется его поведением, а не явным указанием типа.  

🦆 Принцип объясняется поговоркой:  
> *"Если что-то выглядит как утка, плавает как утка и крякает как утка, значит, это утка."*  

В контексте Julia это означает, что если объект поддерживает нужные операции (например, его можно сложить, вызвать метод `push!` и т. д.), то он будет корректно работать в данной функции, независимо от его типа.

<br>

#### 📌 Почему в Julia используется утиная типизация?

- Julia — язык с динамической типизацией по умолчанию. Это значит, что переменные не обязаны иметь строго заданный тип, но могут. 
- Это делает код более гибким и удобным для использования с различными типами данных.
- Улучшает поддержку обобщённого программирования (Generic Programming).
  
  <br>


### **📌 Синтаксис утиных типов в Julia**

#### **1️. Функции без ограничений типов (чистый Duck-typing)**  
В Julia по умолчанию функции принимают параметры любого типа, если они поддерживают нужные операции.

```julia
function add(x, y)
    return x + y  # Оператор `+` работает с числами, строками, векторами
end
```

Output:
```
add (generic function with 1 method)
```

```julia
println(add(3, 5))                 # Выведет 8
println(add("Hello, ", "World!"))  # Выведет "Hello, World!"
println(add([1, 2], [3, 4]))       # Выведет  [1, 2, 3, 4]
```

Output:
```
8
   @ Main ./In[6]:2
 [2] top-level scope
   @ In[7]:2
```

✅ Функция работает с разными типами (`Int`, `String`, `Vector{Int}`), потому что оператор `+` определён для них.

<br>


#### **2️. Использование аннотаций типов (Type Annotations)**  

Если мы хотим ограничить типы передаваемых значений Аргументов, можно использовать `::Тип`.

```julia
function add_numbers(x::Number, y::Number) # Типы аргументов
    return x + y # Оператор `+` работает только с числами (Number) в этом случае
end

println(add_numbers(2, 3))       # Выведет 5
println(add_numbers(3.5, 2.1))   # Выведет 5.6
println(add_numbers(3, 2.1))     # Выведет 5.1
```

Output:
```
5
5.6
5.1
```

✅ Разрешены любые подтипы `Number` (целые, дробные).  

<br>

🔴 Но если передадим строки как Аргументы в нашу функцию, то получим ошибку:

```julia
println(add_numbers("Hello", "World"))  # Ошибка!
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[9]:1
```

**Пример:**

Ограничим значения принимаемых аргументов,  только целым типом чисел `::Int`

```julia
function add_numbers_Int(x::Int, y::Int)
    return x + y
end
```

Output:
```
add_numbers_Int (generic function with 1 method)
```

🔴 Попытка передать `Float64` приведет к ошибке:

```julia
add_numbers_Int(3.5, 2.1)  # Ошибка!
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[11]:1
```

✅ **Но можно разрешить несколько типов:**

```julia
function add_numbers_mix(x::Int64, y::Float64)
    return x + y
end
```

Output:
```
add_numbers_mix (generic function with 1 method)
```

```julia
add_numbers_mix(7, 2.12)
```

Output:
```
9.120000000000001
```

<br>

💡💡💡 **ВНИМАНИЕ !**

**В Julia нет автоматического преобразования типов в сигнатуре функции**

Это означает, что Julia не будет автоматически приводить `Int64` к `Int32`, если в сигнатуре указано `Int32`

**Пример:**

```julia
function add_numbers_mix2(x::Int32, y::Float64)
    return x + y
end
```

Output:
```
add_numbers_mix2 (generic function with 1 method)
```

```julia
add_numbers_mix2(7, 2.12) # Ошибка! 
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[15]:1
```

<br>

**📌 Как же это мы можем это исправить?**

Если мы не хотим ослабить ограничения типов, что означало бы позволение использовать любые числа, указав `::Number` вместо `::Int32`, и хотим, чтобы аргументы были именно `Int32`, то необходимо явно приводить `7` к `Int32` при вызове функции `add_numbers_mix2()`:

```julia
add_numbers_mix2(Int32(7), 2.12) # Вернет 9.12
```

Output:
```
9.120000000000001
```

##### **😺 Выводы:**

- Ошибка происходит, потому что `7` — это `Int64`, а не `Int32`, и Julia не делает автоматическое приведение типов в сигнатуре функции.
   
- Можно ослабить ограничения типов (`Number`), явно приводить аргументы (`Int32(7)`) или использовать `where` (этот способ рассмотрен в этом уроке ниже).
   
- В Julia функции **не перегружаются автоматически для совместимых типов** — необходимо либо указывать универсальные типы, либо явно приводить их. 

**💡 В Julia можно явно ограничивать типы, но часто это не требуется из-за утиных типов.**

<br>



#### **3. Полиморфизм с `Any`**  
Если хотим разрешить *любой* тип данных, можно использовать `::Any`.


```julia
function print_type(x::Any)
    println("Type of x: ", typeof(x)) # typeof(x) возвращает тип переменной x
end

print_type(42)          # Type of x: Int64
print_type("Julia")     # Type of x: String
print_type([1, 2, 3])   # Type of x: Vector{Int64}
```

Output:
```
Type of x: Int64
Type of x: String
Type of x: Vector{Int64}
```

✅ `Any` позволяет передавать абсолютно любой объект.

<br>

#### **4. Использование `where` (рус."Где") для параметрического полиморфизма**  

Можно объявлять функции, принимающие параметры определённых семейств типов.


```julia
function multiply_by_two(x::T) where T <: Number # Где T - это подтип Number
    return x * 2
end

println(multiply_by_two(4))    # 8
println(multiply_by_two(2.5))  # 5.0

println(multiply_by_two("Hi Boy!")) # Ошибка, строка не является Number
```

Output:
```
8
5.0
Stacktrace:
 [1] top-level scope
   @ In[18]:8
```

✅ `T <: Number` означает, что `T` должен быть числом.


<br>

---

<br>

### **3️⃣ Изменяющие (Mutating) и неизменяющие (Non-mutating) функции**

В Julia принято соглашение:  

- **Неизменяющие функции** - это функции, которые ***не изменяют*** переданные им аргументы и обычно записываются без `!` в названии.
  
- **Изменяющие функции** - это функции, которые ***изменяют*** переданный им объект и обычно обозначаются `!` в конце имени.

<br>


##### **1️. Неизменяющая функция `имя()`(создает новый массив):**

```julia
function double_values(arr)
    return arr .* 2
end

a = [1, 2, 3]
b = double_values(a)
println(a)  # [1, 2, 3] (исходный массив не изменился)
println(b)  # [2, 4, 6] (новый массив)
```

Output:
```
[1, 2, 3]
[2, 4, 6]
```

##### **2. Изменяющая функция `имя!()` (модифицирует объект на месте  и обычно обозначаются `!` в конце имени):**

```julia
function double_values!(arr)
    for i in eachindex(arr)
        arr[i] *= 2
    end
end

a = [1, 2, 3]
double_values!(a)
println(a)  # [2, 4, 6] (исходный массив изменился)
```

Output:
```
[2, 4, 6]
```


✅ Изменяющие функции полезны для работы с большими массивами, так как они избегают создания новых копий данных.

<br>

---

<br>

### **4️⃣ Функции высшего порядка (Higher Order Functions)**

Функции высшего порядка принимают другие функции в качестве аргументов или возвращают их.

#### **1. Функция `map`.**

Применяет функцию к каждому элементу коллекции:

```julia
squares = map(x -> x^2, [1, 2, 3, 4])
println(squares)  # [1, 4, 9, 16]
```

Output:
```
[1, 4, 9, 16]
```

#### **2. Функция `filter`.**

Фильтрует элементы, оставляя только те, для которых функция возвращает `true`:

```julia
evens = filter(x -> x % 2 == 0, [1, 2, 3, 4, 5])
println(evens)  # [2, 4]
```

Output:
```
[2, 4]
```

### **3. Функция `reduce`.**

Сворачивает коллекцию с использованием бинарной функции:

```julia
sum_all = reduce(+, [1, 2, 3, 4, 5])
println(sum_all)  # 15
```

Output:
```
15
```

### **4. Функции как аргументы.**



```julia
function apply_twice(f, x)
    return f(f(x))
end

println(apply_twice(x -> x + 1, 5))  # 7

```

Output:
```
7
```

<br>

#### **😺 Выводы:**

- В Julia можно объявлять функции разными способами.
  
- Поддерживается утиная типизация, но можно явно указывать типы.
  
- Изменяющие функции помечаются `!`, неизменяющие — нет.
  
- Функции высшего порядка позволяют передавать другие функции как аргументы.
  
  <br>

**Это делает Julia мощным инструментом для функционального программирования и численных вычислений. 🚀**

<br>

---


<br>


## 🔹 **Упражнения:**

#### ✅ Задание 5.1

Напишите свою функцию `add_one`, которая добавляет 1 к своим входным данным.

```julia
# Ваше решение:


```

<!-- grade_id: cell-5119b9e9623c1cb7, points: 1, grade: true, solution: false -->
```julia
@assert add_one(1) == 2 # Проверка на равенство
```

<!-- grade_id: cell-50f83d27187a2064, points: 1, grade: true, solution: false -->
```julia
@assert add_one(11) == 12 # Проверка на равенство
```

**Пример (правильного решения):**

```julia
function add_one(x)
    return x + 1
end

# Проверки
@assert add_one(1) == 2  # Проверка на равенство
@assert add_one(11) == 12 # Проверка на равенство

println("All tests passed!")  # Если нет ошибок, значит, тесты пройдены

```

Output:
```
All tests passed!
```

<br>

#### ✅ Задание 5.2 

Используйте `map` или `broadcast`, чтобы увеличить каждый элемент матрицы `A` на `1` и присвоить его переменной `A1`.

```julia
# Ваше решение:


```

```julia
@assert A1 == [2 3 4; 5 6 7; 8 9 10] # Проверка на равенство
```

**Примеры (правильного решения):**

1. Вот реализация в Julia, используя `broadcast` и  `.` -  точечный оператор:

```julia
# Создадим матрицу A
A = [1 2 3; 4 5 6; 7 8 9]

# Вариант 1: Использование broadcast (.) 
A1 = A .+ 1  # Каждый элемент матрицы A увеличивается на 1

# Проверка условия
@assert A1 == [2 3 4; 5 6 7; 8 9 10]

println("All tests passed!")  # Если ошибок нет, значит, всё работает
```

Output:
```
All tests passed!
```

📌 Разбор этого кода:

  Использование `broadcast (A .+ 1)`

   -  Оператор `.+` добавляет `1` ко всем элементам матрицы `A`.
  
   - `broadcast` позволяет применять операции поэлементно, работая с массивами любого размера.

2. Альтернативный вариант это использование `map`:

```julia
A1 = map(x -> x + 1, A)  # Применяем функцию ко всем элементам

@assert A1 == [2 3 4; 5 6 7; 8 9 10]
println("All tests passed!")
```

Output:
```
All tests passed!
```

📌 Разбор этого кода

Использование `map` - `(map(x -> x + 1, A))`

- map применяет анонимную функцию `x -> x + 1` ко всем элементам матрицы `A`.

- Возвращает новый массив с результатами.

<br>


#### ✅ Задание 5.3 

Используйте точечный синтаксис `broadcast` для увеличения каждого элемента матрицы `A1` на 1 и сохраните результат в переменной `A2`.

Затем убедитесь, что `A2 == [3 4 5; 6 7 8; 9 10 11]` с помощью `@assert`.

Ваше решение:

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
# Ваше решение:


```

**Пример (правильного решения):**

```julia
# Исходная матрица A1
A1 = [2 3 4; 5 6 7; 8 9 10]

# Используем broadcast (точечный оператор .+) для увеличения каждого элемента на 1
A2 = A1 .+ 1

# Проверка условия
@assert A2 == [3 4 5; 6 7 8; 9 10 11]

println("All tests passed!")  # Если ошибок нет, значит всё работает правильно
```

Output:
```
All tests passed!
```

✅ Решение использует эффективный и idiomatic (естественный для Julia) способ обработки массивов. 🚀